# Part 1: An introduction to jupyter notebooks and python

We are going to start with some basic aspects of the jupyter environment. I will give a brief lecture about the history of the notebook concept and some basic python commands. We will go through each of the menu options above and review the key elements of the notebook environment.

**The learning objectives for this exercise are:**
1. Understand how to assign a value to a variable
2. Understand how to add, multiply, and divide variables
3. Understand how to compare values (<,>, AND, NOT, OR)
3. Understand that variables have a type
4. Learn to write a *function*

**Resources for learning python, and getting more help:**
- Numpy Quickstart Tutorial: https://docs.scipy.org/doc/numpy/user/quickstart.html
- Markdown: [Markdown reference](https://jupyter-notebook.readthedocs.io/en/stable/examples/Notebook/Working%20With%20Markdown%20Cells.html)
- CodeAcademy Python: https://www.codecademy.com/catalog/language/python


# Assigning values to variables

The most basic operation in any programming language is *assignment*. When we assign a value to a *variable*, we are telling the programming language that we want to be able to refer to the stored value using the variable's name. We will start by assigning the number 2 to a variable called 'x'. Then, we can ask python to print out the value of x. We do this by typing 'print(x)'.

<!-- grade_id: assignment, points: 0, grade: false, solution: true -->
```python
x=2
print(x)
```

<!-- grade_id: cell-104492d488bd48ee, points: 1, grade: true, solution: false -->
```python
# This is an autograder test. For now, you can ignore the code within this cell.
assert( x==2 )
```

# Arithmetic
We can now use python programming to manipulate the variable `x`. Let's calculate the values of $x^2$ and $\sqrt{x}$ and print them out. We can evalute $a^b$ using the code `a**b`.

<!-- grade_id: cell-6f40127c5fc4b515, points: 0, grade: false, solution: true -->
```python
x**0.5
```

<!-- grade_id: cell-e5f26b252d2cd826, points: 0, grade: false, solution: true -->
```python
raise NotImplementedError()
```

We can use arithmetic operations to add, multiply, and divide variables. The *operators* for these operations are listed below:

| Symbol | Operation |
|--------|-----------|
| +	| Addition |
| -	|  Subtraction |
| /	|  division |
|  \% |	mod |
| *	| multiplication |
| // | floor division |
| \*\* | to the power of |


In the cell below, take some time to try out each of the arithmetic operations.

<!-- grade_id: cell-3379d2e4d26d4ff8, points: 1, grade: true, solution: true -->
```python
11 % 4
```

# Comparisons and logic
We will frequently want to compare the value of one variable with another, or see whether two things are equal. To this end, Python provides logical operators. The comparison operators are listed in the table below:

| Symbol | Operation |
|--------|-----------|
| a > b  | Greater than |
| a >= b | Greater than or equal to |
| a <  b | Less than |
| a <= b | Less than or equal to |
| a == b | equal to |
| a != b | not equal to |

**Note - one of the most common beginner mistakes in coding is to confuse the = and == operators. '=' performs assignment whereas '==' performs a comparison.**

When you compare two quantities, the result is always True or False. True and False are examples of Logical values. Note that True and False are capitalized and that case matters! Use the cell below to experiment with logical comparisons.

<!-- grade_id: cell-c11c3df9991fbc82, points: 1, grade: true, solution: true -->
```python
a=5
b=5
a != b
```

# Types

In python and other programming language, variables store a specific *type* of data. For now, there are three important data types that we have worked with:

- Integers ('int type')
- Floating point numbers ('float')
- Logical
- Strings (text characters)

Often we don't need to think about data types, but sometimes they become important. For instance, sometimes we need to represent data as an integer, or as a floating point number. When we make comparisons between data types, we often need to be careful. For instance, while 2.0 and 2 mean the same thing to us when working with paper and pencil, they may be represented as a floating point number and an integer, respectively.

If you need to know the data type of a variable (and sometimes this is *very* useful for de-bugging code), you can ask python to print out the type of the variable, as shown below:

```python
a=2
print("The type of a is " + str(type(a))) # note that str(x) converts x into a string
b=2.0
print("The type of a is " + str(type(b)))
```

The code above demonstrates something important, which is that by default, a number is represented as an integer if it's a whole number. However, you can force python to represent a number as a floating point number if you include a decimal point (either 'b=2. or b=2.0')

# Functions 
One of the most common and useful things to do in any programming language is to create a function to carry out a task. The reason that we do this is that often we want to apply the same procedure repeatedly to many different datasets. For instance, we might want to calculate the mean of many different populations.

For our first function, let's write  a function to  take the square-root of an argument. To do this, we have to define the function by giving it a name, which in this case will be `sqrt`. A function operates on some *input* called the *argument* and returns one or more outputs, called the *return value*. We will call the argument `x`. The syntax for our function will be:

```def sqrt(x):
    return x**0.5
    ```
    
Note that the second line, following the `:` is indented. This is very important! In python, *blocks* of code are organized based on indentation level.

<!-- grade_id: cell-392828d0c890ba1c, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: correct_squares, points: 1, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
assert sqrt(4) == 2.0
```

# Flow control

Flow control refers to the use of logical statements to send your program down different 'branches' of execution. In mathematics, we commonly run into functions that have *piecewise* definitions. See the example below:

$$
f(x) =  \begin{cases} 
      0 & x\leq 0 \\
      \frac{50-x}{50} & 0\leq x\leq 50 \\
      0 & 50\leq x 
   \end{cases}
$$

If we wanted to implement a python function to calculate the values of $f(x)$, we could do the following:

```python
def f(x):
    if x<=1.0:
        return 0.0
    elif x<=50:
        return (50.-x)/50.
    else:
        return 0.0
```

```python
f(49.99)
```

## On your own
Write a function that calculates the absolute value of x (without using the built-in absolute value function).

```python

```