# Part 2: Introduction to NumPy and a first plot

In Part 1, we looked at some basic elements of the python programming language. Now we're going to look at some additional features of the language and introduce a powerful 'package' for carrying out computational tasks in python. The objectives of this part of the exercise are:

1. Introduce the NumPy package and learn how to access it.
2. Introduce the NumPy array data type.
3. Use a loop to perform an operation on data stored in an array.
4. Make a simple graph (plot)

# NumPy

Numpy is an example of a python *package*, which is a collection of functions and data structures that can be loaded and used in your programs. Almost everything that we will do in this course will make use of NumPy in some way. The core functionality in numpy is the introduction of an *array* data type, which can hold a list of floating-point values. NumPy contains a variety of different functions to efficiently access and manipulate data stored in arrays.

To load numpy, we use the special python command 'import':

'''import numpy as np'''

If at any time you need help with NumPy, you can access the help by clicking the **Help** menu above and then clicking **NumPy Reference**

```python
#This will import all of the numpy functions and assign them the prefix 'np'
import numpy as np
```

As a first example, let's create a numpy array containing a list of numbers between 0-10.

```python
# Note - the data enclosed by square brackets [] are a list of numbers.
x = np.array([1,2,3,4,5,6,7,8,9,10])
print(x)
```

As you can see, x contains the numbers between 1-10. The array has some important properties that we can examine. First, its shape:

```python
print(x.shape)
```

One-dimensional arrays in NumPy also have a well-defined property called 'size'. Print out the size of x below:

<!-- grade_id: cell-10f2635f881fd17e, points: 1, grade: true, solution: true -->
```python
x.size
```

As we can see above, the shape of x is (10,). This means that the array has only one dimension and that it's 10 elements long. NumPy can store multidimensional data (like matrices or 3D volume data) as 2- or 3- dimensional arrays, which we will examine later.

Now, what can we do with the data once it's stored in an array? One thing that we might want to do is calculate the mean of the values in the array and assign the mean to a new variable called x_mean. NumPy has some functionality that allows us to calculate the mean with a single command.

```python
x_mean = x.mean()
print(x_mean)
```

The command above is a little bit different from those that we've seen before. Taking apart the syntax:
- 'x' tells python that we want to do something with the variable 'x'.
- The '.' after x tells python that we want to access some property of *method* of 'x'. A method is a function that runs on 'x'. A NumPy array has many different methods built in that can be used to manipulate its values.
- 'mean()' tells python that we want to calculate the mean of the values in x.

To see a full list of the built in methods that can act on a numpy array, go to Help->NumPy Reference and then click on 'ndarray'.

# Task:
Use the NumPy documentation for 'ndarray' (see above) to find out how to calculate the sum of all of the elements in the array x. Assign the result to the variable 'x_sum':

<!-- grade_id: cell-39e5b34d2fe704bc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a3f7cd1fe66403b8, points: 1, grade: true, solution: false -->
```python
assert(x_sum == 55)
```

# Loops

Loops allow us to perform the same task repeatedly on changing inputs. There are a couple of ways of looping over things in python. One of the most convenient ways is shown below:

``for value in x:
        print( value )
``

Copy this code below and see what it does:

<!-- grade_id: cell-9c475608cc25e96c, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

If your code ran correctly, you should see that python prints out each value in the array x. Now, try to do something a little more complicated. Modify the previous 'for' loop so that it prints out the square of each of the numbers:

<!-- grade_id: cell-45b966c848fb8155, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Sometimes it is useful to read or modify the value at a specific location in an array. In python, we access the value at a specific location in an array using its *index*. In python, the first item in an array has the index '0' and the last item has an index of 'N-1' where 'N' is the length of the array.

We can access a specific item in x using the syntax:
``x[index]``

For example, if we want the second item, we can type:
``x[1]``
(remember: the second item has an index of 1, because the first item has an index of 0).
Use the cell below to experiment with accessing elements in x. Try accessing an element with an index greater than (N-1) or less than 0 and think about what happens.

<!-- grade_id: cell-c49f9b6d99c6105d, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
for i in range(0,x.size):
    print( x[i]**2 )
```

# Your first plot

There is a python package called matplotlib that contains functions for preparing figures and visualizing data. Let's go through an example of how to plot the values of a function and prepare a neatly labeled figure:

```python
import matplotlib.pyplot as plt
plt.figure()                        # This creates a new figure window
t = np.linspace(-1.0,1.0,200)       # Generate 200 uniformly spaced values between -1.0 and 1.0
y = np.cos(np.pi*t)                 # Calculate cos(pi*t)
plt.plot(t,y,label='cos($\pi t$)')  # plot t on the horizontal axis and y on the vertical axis. label the curve.
plt.legend()                        # add a legend. the label in the legend is taken from the previous 'label'
plt.xlabel('x')
plt.ylabel('y')
plt.show()                          # show the plot
```

# Task:

Make a plot of the function $f(x)=x^2$ over the interval $-2\le x \le 2$. Use a black line to plot the function. Label the horizontal axis 'x' and the vertical axis 'f(x)'. Add a legend with the formula for the function.

*Hint:* To change the line style of the curve, see the documentation for 'plot' by selecting **Help->Matplotlib Reference** above. Then click on Tutorials and look at the 'pyplot' tutorial

<!-- grade_id: cell-009e7b4f3cb5cb94, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# A word of caution on copying NumPy arrays
# This is important stuff!!!
When we store information in an array, the actual numbers are encoded in a region of the computer's memory. To a large extent, Python tries to hide memory management from us so that we can focus on higher level concepts. However, sometimes we must stop and think about what happens when we assign something to a variable. 
- When we create a new numpy array (e.g. using np.zeros(10,1), a region of memory is allocated to hold the new array and the values are all set to 0.0.
- When we assign an existing array to a new variable, the default behavior is that we are assigning a new name to existing information. The information is not copied, we just get a new 'view' into the existing information. **If you are not careful, this can get you into trouble**. If you are familiar with C/C++, this behavior is equivalent to creating a new pointer to an existing array.
- If, instead, you want to make a copy of the array, so as to preserve the contents of the original array, you can do so using the `copy()` method built into the ndarray.

You can read more about the behavior here [https://docs.scipy.org/doc/numpy/user/quickstart.html#no-copy-at-all]

The behavior is illustrated below:

```python
# create a NumPy array called a
a = np.linspace(0,1,4)
print("a=",a)
b=a     # assign the name 'b' to the existing array 'a'
b[1]=10 # modify the contents of b
print("a=",a) # print out a. Note that by modifying b, we modified a. That's because b is just another name for a
```

```python
# create a NumPy array called a
a = np.linspace(0,1,4)
print("a=",a)
b=a.copy()     # copy the contents of a into a new array and give it the name 'b'
b[1]=10        # modify the contents of b
print("a=",a)  # print out a. Note that a is unchanged
print("b=",b)  # print out b. Note that b has been modified
```

```python

```